```python
# Initialize Otter
import otter
grader = otter.Notebook("hw05.ipynb")
```

<img src="data8logo.png" alt="Data 8 Logo" style="width: 15%; float: right; padding: 1%; margin-right: 2%;"/>

# Homework 5: Applying Functions and Iteration

Please complete this notebook by filling in the cells provided. Before you begin, execute the previous cell to load the provided tests.

**Helpful Resource:**
- [Python Reference](http://data8.org/fa24/reference/): Cheat sheet of helpful array & table methods used in Data 8!

**Recommended Readings**: 

* [Tabular Thinking Guide](https://drive.google.com/file/d/1D9PXUgDLXjFBXZoyGVSKpzk0nKQziDTP/view?usp=sharing)
* [Applying Functions](https://www.inferentialthinking.com/chapters/08/1/Applying_a_Function_to_a_Column.html)
* [Conditionals](https://www.inferentialthinking.com/chapters/09/1/Conditional_Statements.html)
* [Iteration](https://www.inferentialthinking.com/chapters/09/2/Iteration.html)

Please complete this notebook by filling in the cells provided. **Before you begin, execute the cell below to setup the notebook by importing some helpful libraries.** Each time you start your server, you will need to execute this cell again.

For all problems that you must write explanations and sentences for, you **must** provide your answer in the designated space. Moreover, throughout this homework and all future ones, **please be sure to not re-assign variables throughout the notebook!** For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!

**Deadline:**

This assignment is **due Wednesday, 10/2 at 5:00pm PT**. Submissions after this time will be accepted for 24 hours and will incur a 20% penalty. Any submissions later than this 24 hour period will not be accepted unless an extension has been granted as per the [policies](http://data8.org/fa24/policies/) page. Turn it in by Tuesday, 10/1 at 5:00pm PT for 5 extra credit points.

<span style="color:red">**Note: This homework has hidden tests on it. That means even though tests may say 100% passed, it doesn't mean your final grade will be 100%. We will be running more tests for correctness once everyone turns in the homework.**</span>

Directly sharing answers is not okay, but discussing problems with the course staff or with other students is encouraged. Refer to the [policies](http://data8.org/fa24/policies/#learning-cooperatively) page to learn more about how to learn cooperatively.

You should start early so that you have time to get help if you're stuck. Office hours are held Monday through Friday in [Warren Hall](https://www.berkeley.edu/map?warren) 101B. The office hours schedule appears [here](http://data8.org/fa24/schedule/).

---

The point breakdown for this assignment is given in the table below:
| Category | Points |
| --- | --- |
| Autograder (Coding questions) | 100 |
| Written | 0 |
| **Total** | 100 |

<span style="color:red">**Note:**</span> Since there are no written/manual questions in this homework, there is no PDF to submit.

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 0. Midterm Accommodations Form

---

**Question 1.** The DATA C8 Summer 2024 Midterm Exam will take place on **Friday, July 12th from 11AM - 1PM PT**. Please complete this form so that we can best accommodate you for the midterm.
_**All students are required to fill out this form**_. The deadline to submit this form is **Monday, July 8th by 11:59PM.** The link can be found below. **(1 Point)**

- [Fall 2024 Midterm Accommodations Survey](https://docs.google.com/forms/d/e/1FAIpQLSe-O5m6xPtjh-c4vwxoKwU0Ip8uzbM8oN4qPY2rz4HIwNhoaQ/viewform?usp=sf_link)

Assign `secret_phrase` to the secret phrase given at the end of the accommodations survey. Make sure the phrase is in quotes (i.e. is a string)!


```python
secret_phrase = "Can you show me the exit sign"
```

```python
grader.check("q0_1")
```

Output:
```
q0_1 results: All test cases passed!
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 1. 2023 Cal Football Season

```python
# Run this cell to set up the notebook, but please don't change it.

# These lines import the Numpy and Datascience modules.
import numpy as np
from datascience import *

# These lines do some fancy plotting magic.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore', FutureWarning)
```

Output:
```
/Users/thanglee/something beautiful/uni_datascience/.venv/lib/python3.9/site-packages/urllib3/__init__.py:35: NotOpenSSLWarning: urllib3 v2 only supports OpenSSL 1.1.1+, currently the 'ssl' module is compiled with 'LibreSSL 2.8.3'. See: https://github.com/urllib3/urllib3/issues/3020
  warnings.warn(
```

Jonathan and Jadyn are trying to analyze how well the Cal football team performed in the 2023 season. A football game is divided into four periods, called quarters. The number of points Cal scored in each quarter and the number of points their opponent scored in each quarter are stored in a table called `cal_fb.csv`.

```python
# Just run this cell
# Read in the cal_fb csv file
games = Table().read_table("cal_fb.csv")
games.show()
```

Output:
```
<IPython.core.display.HTML object>
```

Let's start by finding the total points each team scored in a game.

---

**Question 1.** Write a function called `sum_scores`.  It should take four arguments, where each argument represents integers corresponding to the team's score for each quarter. It should return the team's total score for that game as an integer. **(2 Points)**

*Hint:* Don't overthink this question!



```python
def sum_scores(one, two, three, four):
    return one + two + three + four

sum_scores(14, 7, 3, 0) #DO NOT CHANGE THIS LINE
```

Output:
```
24
```

```python
grader.check("q1_1")
```

Output:
```
q1_1 results: All test cases passed!
```

---

**Question 2.** Create a new table `final_scores` with three columns in this *specific* order: `Opponent`, `Cal Score`, `Opponent Score`. You will have to create the `Cal Score` and `Opponent Score` columns. Use the function `sum_scores` you just defined in the previous question for this problem. **(5 Points)**

*Hint:* If you want to apply a function that takes in multiple arguments, you can pass multiple column names as arguments in `tbl.apply()`. The column values will be passed into the corresponding arguments of the function. Take a look at the Python Reference Sheet and Lecture 13's demo for syntax.

*Note:* If you’re running into issues creating `final_scores`, check that `cal_scores` and `opp_scores` output what you want. If you're encountering `TypeError`s, check the [Python Reference](https://www.data8.org/fa24/reference/) to see if the inputs/outputs of the function are what you expect.


```python
cal_scores = games.apply(sum_scores, "Cal 1Q", "Cal 2Q", "Cal 3Q", "Cal 4Q")

opp_scores = games.apply(sum_scores, "Opp 1Q", "Opp 2Q", "Opp 3Q", "Opp 4Q")

final_scores = Table().with_columns(
    "Opponent", games.column("Opponent"),
    "Cal Score", cal_scores,
    "Opponent Score", opp_scores
)
final_scores
```

Output:
```
Opponent         | Cal Score | Opponent Score
North Texas      | 58        | 21
Auburn           | 10        | 14
Idaho            | 31        | 17
Washington       | 32        | 59
Arizona State    | 24        | 21
Oregon State     | 40        | 52
Utah             | 14        | 34
USC              | 49        | 50
Oregon           | 19        | 63
Washington State | 42        | 39
... (3 rows omitted)
```

```python
grader.check("q1_2")
```

Output:
```
q1_2 results: All test cases passed!
```

We can get specific row objects from a table. You can use `tbl.row(n)` to get the `n`th row of a table. `row.item("column_name")` will allow you to select the element that corresponds to `column_name` in a particular row. Here's an example:

```python
# Just run this cell
# We got the Axe!
games.row(10) # <-- this will return a row object
```

```python
# Just run this cell
games.row(10).item("Cal 4Q") # <-- this will return a item (e.g. an int) from a row object
```

---

**Question 3.** We want to see for a particular game whether or not Cal lost. Write a function called `did_cal_lose`.  It should take one argument: a **row object** from the `final_scores` table. It should return either `True` if Cal's score was less than the Opponent's score, and `False` otherwise. **(5 Points)**

*Note 1*: "Row object" means a row from the table extracted (behind the scenes) using `tbl.row(index)` that contains all the data for that specific row. It is **not** the index of a row. Do not try and call `final_scores.row(row)` inside of the function.

*Note 2*: If you're still confused by row objects, try printing out `final_scores.row(1)` in a new cell to visually see what it looks like! This piece of code is pulling out the row object located at index 1 of the `final_scores` table and returning it. When you display it in a cell, you'll see that it is not located within a table, but is instead a standalone row object!

Hint: You can use comparators like <, >, <=, >=, ==, and != to compare items from a row object


```python
def did_cal_lose(row):
    return row.item("Cal Score") < row.item("Opponent Score")

did_cal_lose(final_scores.row(1)) #DO NOT CHANGE THIS LINE
```

Output:
```
True
```

```python
grader.check("q1_3")
```

Output:
```
q1_3 results: All test cases passed!
```

---

**Question 4.** Jonathan and Jadyn wants to see how Cal did against every opponent during the 2023 season. Using the `final_scores` table:

1. Assign `results_array` to an array of `True` and `False` boolean values that correspond to whether or not Cal lost.
2. Add `results_array` to the `final_scores` table in a column named `Results`, and assign this to `final_scores_with_results`.
3. Then, respectively assign the number of wins and losses Cal had to `cal_wins` and `cal_losses`.

**(10 Points)**

*Hint*: `tbl.apply()` might be helpful. Refer to the [Python Reference](https://www.data8.org/fa24/reference/) if you're unsure how it works!


```python
results_array = final_scores.apply(did_cal_lose)
final_scores_with_results = final_scores.with_columns("Results", results_array)
cal_losses = final_scores_with_results.where("Results", are.equal_to(True)).num_rows
cal_wins = final_scores_with_results.where("Results", are.equal_to(False)).num_rows

# Don't delete or edit the following line:
print(f"In the 2023 Season, Cal Football won {cal_wins} games and lost {cal_losses} games. Go Bears! 🐻")
```

Output:
```
In the 2023 Season, Cal Football won 6 games and lost 7 games. Go Bears! 🐻
```

```python
grader.check("q1_4")
```

Output:
```
q1_4 results: All test cases passed!
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 2. Unrolling Loops (Optional)

**This section of HW5 is optional. Do it for your own practice, but it will not be incorporated into the final grading!**

"Unrolling" a `for` loop means to manually write out all the code that it executes.  The result is code that does the same thing as the loop, but without the structure of the loop.  For example, for the following loop:

    for num in np.arange(3):
        print("The number is", num)

The unrolled version would look like this:

    print("The number is", 0)
    print("The number is", 1)
    print("The number is", 2)


Unrolling a `for` loop is a great way to understand what the loop is doing during each step. In this exercise, you'll practice unrolling a `for` loop.


In the question below, write code that does the same thing as the given code, but with any `for` loops unrolled.  It's a good idea to run both your answer and the original code to verify that they do the same thing.  (Of course, if the code does something random, you'll get a different random outcome than the original code!)

---

**Optional Question 1.** Unroll the code below.

```python
for joke_iteration in np.arange(3):
    print("Knock. " * (joke_iteration + 1))
    print("Who's there?")
    print("Banana.")
    print("Banana who?")
print("Knock, knock.")
print("Who's there?")
print("Orange.")
print("Orange who?")
print("Orange you glad I didn't say banana?")
```

Output:
```
Knock. 
Who's there?
Banana.
Banana who?
Knock. Knock. 
Who's there?
Banana.
Banana who?
Knock. Knock. Knock. 
Who's there?
Banana.
Banana who?
Knock, knock.
Who's there?
Orange.
Orange who?
Orange you glad I didn't say banana?
```

```python
# First iteration (joke_iteration = 0)
print("Knock. ")
print("Who's there?")
print("Banana.")
print("Banana who?")

# Second iteration (joke_iteration = 1)
print("Knock. Knock. ")
print("Who's there?")
print("Banana.")
print("Banana who?")

# Third iteration (joke_iteration = 2)
print("Knock. Knock. Knock. ")
print("Who's there?")
print("Banana.")
print("Banana who?")

# After the loop finishes
print("Knock, knock.")
print("Who's there?")
print("Orange.")
print("Orange who?")
print("Orange you glad I didn't say banana?")

```

Output:
```
Knock. 
Who's there?
Banana.
Banana who?
Knock. Knock. 
Who's there?
Banana.
Banana who?
Knock. Knock. Knock. 
Who's there?
Banana.
Banana who?
Knock, knock.
Who's there?
Orange.
Orange who?
Orange you glad I didn't say banana?
```

Congratulations, you are done with Homework 5!

## Pets of Data 8
**Tito** says congrats on finishing Lab 6!

<img src="./tito.jpg" width="40%" alt="Black and white cat sitting next to window"/>

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Output:
```
Running your submission against local test cases...


Your submission received the following results when run against available test cases:

    q0_1 results: All test cases passed!

    q1_1 results: All test cases passed!

    q1_2 results: All test cases passed!

    q1_3 results: All test cases passed!

    q1_4 results: All test cases passed!
<IPython.core.display.HTML object>
```

Tests for q0_1 (points: None)
```
>>> # Make sure that your answer is a string
>>> type(secret_phrase) == str 
True
```


Tests for q1_1 (points: [1, 1])
```
>>> sum_scores(2, 3, 6, 1)
12
```
```
>>> sum_scores(-2,3,5,-10)
-4
```


Tests for q1_2 (points: [0, 0, 0, 0, 0])
```
>>> final_scores.num_columns
3
```
```
>>> final_scores.num_rows
13
```
```
>>> set(['Opponent', 'Cal Score', 'Opponent Score']) == set(final_scores.labels)
True
```
```
>>> isinstance(final_scores.column(0).item(0), str)
True
```
```
>>> isinstance(final_scores.column(1).item(0), int)
True
```


Tests for q1_3 (points: [0, 2])
```
>>> 'did_cal_lose' in globals()
True
```
```
>>> did_cal_lose(final_scores.row(1))
True
```


Tests for q1_4 (points: [0, 0, 0])
```
>>> 0 <= cal_wins <= 13
True
```
```
>>> 0 <= cal_losses <= 13
True
```
```
>>> "Results" in final_scores_with_results.labels
True
```
